<a href='https://ai.meng.duke.edu'> = <img align="left" style="padding-top:10px;" src=https://storage.googleapis.com/aipi_datasets/Duke-AIPI-Logo.png>

# In-class Warmup Exercise
### Question 1
Let's start with a warmup exercise using just Python.  We are going to perform some basic analysis of the stock price for a fictional company NewCo.  Let's suppose we receive some data on the last 10 days of NewCo's stock price.  The data is in the form of a dictionary 'newco_dict' where the keys are ['DATE','LOW','HIGH','CLOSE'].  'DATE' contains a list of string dates. 'LOW' and 'HIGH' each contain a list of the lowest price and highest price respectively of the stock during each day in the 10-day period. 'CLOSE' contains a list of the closing price of the stock at the end of each day during the 10-day period.

For example, newco_dict might look like the following:  

```
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}
```

Our goal in this assignment will be to find all the dates during the 10-day period when the closing price of the stock was equal to the mode (most frequent) value of the closing prices over the 10-day period.  We are going to tackle this in two parts.  First we will create a function to find the mode of the closing prices.  In part 2 we will then create a function to identify the dates on which the closing price was equal to the mode.


NOTE: For all parts of Question 1, you may not import any additional libraries - I would like you to solve these using only the Python you have learned so far and no helper libraries unless already imported for you. 

#### Question 1.1
Let's first define a function `most_freq_closing(price_dict)` which takes as input a dictionary of prices in the form of 'newco_dict' and returns the most frequent closing price (the mode of the closing prices) during the period as an integer.  You may assume that the dictionary is non-empty and that there is a single mode price (in real life, some datasets may have no mode, if all elements are represented equally as often, or may have more than one mode).

There are many ways to approach this problem.  One method (although there are others) might be to iterate through the set of the unique closing values and count how often each one appears during the 10 days.  The mode is the one that appears most frequently (has the highest count of appearances).  

How you structure your solution is up to you, although you must complete the `most_freq_closing` function below.  You may include all your solution code within that function or break sub-tasks out into separate smaller functions (which you can insert into the same cell or add a new cell).

<!-- grade_id: cell-c61d21d175f8cc1a, points: 0, grade: false, solution: true -->
```python
def most_freq_closing(price_dict):
    '''
    Finds the most frequent closing price of a stock

    Inputs:
        price_dict(dict): Dictionary containing DATE, LOW, HIGH, CLOSE prices for an arbitrary number of days

    Returns:
        mode_val(int): Most frequent closing price
    '''
    ### BEGIN SOLUTION ###   #1)create new object --> close price 2) go through list and count the most unique way 
    
    close_price = price_dict['Close']
    for price in close_price
    if price 
     


        
   
    ## END SOLUTION ### 
```

Output:
```
Input In [2]
    if price in price_dict['CLOSE']
                                   ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-1e5d837902352889, points: 1, grade: true, solution: false -->
```python
# `most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}
mode_price = most_freq_closing(newco_dict)
print(f'Your function calculated the most frequent price as {mode_price}')
assert type(mode_price) == int, 'Your function did not return an integer'
assert mode_price == 75,'Incorrect, check your function and try again'

newco_dict2 = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,74,90,86,90,93,100,91,75,68]}
mode_price = most_freq_closing(newco_dict2)
assert type(mode_price) == int, 'Your function did not return an integer'
assert mode_price == 90,'Incorrect, check your function and try again'
print('\n(Passed tests!)')
```

#### Question 1.2
Let's now make use of our new function `most_freq_closing` to find all dates on which the most frequent closing price occurred.  Write a new function `dates_most_freq_closing(price_dict)` which finds and returns the dates on which the most frequent (mode) closing price occurs.  The input to the function is a dictionary of the form of `newco_dict` and the output is a **list** of dates.  Your function should make use of the function `most_freq_closing()` that you just wrote to find the mode closing price, and then you should find all dates on which it occurs.

<!-- grade_id: cell-2f9bc10ff3889fa4, points: 0, grade: false, solution: true -->
```python
def dates_most_freq_closing(price_dict):
    '''
    Finds the dates on which the closing price is the mode of the time period of the data proved as an input

    Inputs:
        price_dict(dict): Dictionary containing DATE, LOW, HIGH, CLOSE prices for an arbitrary number of days

    Returns:
        dates(list): List of dates on which the close price is the mode
    '''
    ### BEGIN SOLUTION ### 
    
    ### END SOLUTION ###
```

<!-- grade_id: cell-2dddb819d9fe493a, points: 1, grade: true, solution: false -->
```python
# `dates_most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}

# Check that dates_most_freq_closing uses most_freq_closing
orig_most_freq_closing = most_freq_closing
del most_freq_closing
try:
    dates_most_freq_closing(newco_dict)
except NameError:
    pass
else:
    raise AssertionError("dates_most_freq_closing does not use most_freq_closing")
finally:
    most_freq_closing = orig_most_freq_closing
    
# Check output
dates = dates_most_freq_closing(newco_dict)
print(f'Your function returned the dates: {dates}')
assert type(dates) == list, 'Your function did not return a list of dates'
assert type(dates[0] == str), 'The values in your list are not strings'
assert dates == ['06-15-21', '06-16-21', '06-24-21'], 'Incorrect, check your function and try again'
print('\n(Passed tests!)')
```

### Question 2
#### Question 2.1
Let's first define a function `most_freq_closing(price_dict)` which takes as input a dictionary of prices in the form of 'newco_dict' and returns the most frequent closing price (the mode of the closing prices) during the period as an integer.  You may assume that the dictionary is non-empty and that there is a single mode price (in real life, some datasets may have no mode, if all elements are represented equally as often, or may have more than one mode).

NOTE: for this problem you MUST use the NumPy library, which I have imported above (make sure you already have NumPy installed - if not first follow the instructions in the setup.md document in the Module 2 Introduction section of Sakai).  **You should not use any loops (including list comprehensions) in your solution.**

One solution might be to count the occurrences of each unique closing value during the period, and the find the closing value with the highest count of occurences.  You can do these steps using NumPy without any loops.  There are multiple ways to solve this.  You may include all your code in the `most_freq_closing` function below or you may choose to add additional functions for each sub-task which are then called by `most_freq_closing`.

<!-- grade_id: cell-bef5a9af5abd959c, points: 0, grade: false, solution: true -->
```python
import numpy as np

def most_freq_closing(price_dict):
    '''
    Finds the most frequent closing price of a stock

    Inputs:
        price_dict(dict): Dictionary containing DATE, LOW, HIGH, CLOSE prices for an arbitrary number of days

    Returns:
        mode_val(int): Most frequent closing price
    '''
    ### BEGIN SOLUTION ###
    



    ### END SOLUTION ###
```

<!-- grade_id: cell-4f25961edf86ff8d, points: 1, grade: true, solution: false -->
```python
# `most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}
mode_price = most_freq_closing(newco_dict)
print(f'Your function calculated the most frequent price as {mode_price}')
assert type(mode_price) == int, 'Your function did not return an integer'
assert mode_price == 75,'Incorrect, check your function and try again'

newco_dict2 = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,90,86,90,93,100,91,75,68]}
mode_price = most_freq_closing(newco_dict2)
print(f'Your function calculated the most frequent price as {mode_price}')
assert type(mode_price) == int, 'Your function did not return an integer'
assert mode_price == 90,'Incorrect, check your function and try again'
print('\n(Passed tests!)')
```

Output:
```
/Users/brunovalan/Class/AIPI_510/aipi510_repo/week3/inclass_exercise/intro_np_pd/warmup.ipynb Cell 10 in most_freq_closing(price_dict)
     <a href='vscode-notebook-cell:/Users/brunovalan/Class/AIPI_510/aipi510_repo/week3/inclass_exercise/intro_np_pd/warmup.ipynb#X12sZmlsZQ%3D%3D?line=13'>14</a> close_price = price_dict['CLOSE']
     <a href='vscode-notebook-cell:/Users/brunovalan/Class/AIPI_510/aipi510_repo/week3/inclass_exercise/intro_np_pd/warmup.ipynb#X12sZmlsZQ%3D%3D?line=14'>15</a> array_close_price = np.array(close_price)
---> <a href='vscode-notebook-cell:/Users/brunovalan/Class/AIPI_510/aipi510_repo/week3/inclass_exercise/intro_np_pd/warmup.ipynb#X12sZmlsZQ%3D%3D?line=15'>16</a> mode_val = np.bincount(array_close_price.argmax())
     <a href='vscode-notebook-cell:/Users/brunovalan/Class/AIPI_510/aipi510_repo/week3/inclass_exercise/intro_np_pd/warmup.ipynb#X12sZmlsZQ%3D%3D?line=16'>17</a> return mode_val

File <__array_function__ internals>:5, in bincount(*args, **kwargs)

Valu
```

#### Question 2.2
Let's now make use of our new function `most_freq_closing` to find all dates on which the most frequent closing price occurred.  Write a new function `dates_most_freq_closing(price_dict)` which finds and returns the dates on which the most frequent (mode) closing price occurs.  The input to the function is a dictionary of the form of `newco_dict` and the output is a **list** of dates.  Your function should make use of the function `most_freq_closing()` that you just wrote to find the mode closing price, and then you should find all dates on which it occurs.

NOTE: for this problem you MUST use the NumPy library, which I have imported above (make sure you already have NumPy installed - if not first follow the instructions in the setup.md document in the Module 2 Introduction section of Sakai).  You MAY NOT use any loops (including list comprehensions) in your solution.

Hint (only read if you need it, otherwise skip!):  
- You might find it helpful to filter an array of the closing prices to those which match the most frequent price, and then use the resulting Boolean mask (see example in demo notebook) to filter the dates

```python
def dates_most_freq_closing(price_dict):
    '''
    Finds the dates on which the closing price is the mode of the time period of the data proved as an input

    Inputs:
        price_dict(dict): Dictionary containing DATE, LOW, HIGH, CLOSE prices for an arbitrary number of days

    Returns:
        dates(list): List of dates on which the close price is the mode
    '''
    ### BEGIN SOLUTION ###
    
    ### END SOLUTION ###
```

```python
# `dates_most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}

# Check that dates_most_freq_closing uses most_freq_closing
orig_most_freq_closing = most_freq_closing
del most_freq_closing
try:
    dates_most_freq_closing(newco_dict)
except NameError:
    pass
else:
    raise AssertionError("dates_most_freq_closing does not use most_freq_closing")
finally:
    most_freq_closing = orig_most_freq_closing
    
# Check output
dates = dates_most_freq_closing(newco_dict)
print(f'Your function returned the dates: {dates}')
assert type(dates) == list, 'Your function did not return a list of dates'
assert type(dates[0] == str), 'The values in your list are not strings'
assert dates == ['06-15-21', '06-16-21', '06-24-21'], 'Incorrect, check your function and try again'
print('\n(Passed tests!)')
```